<div style="color:white; border-radius:8px; background-color:#03112A; font-size:250%; letter-spacing:1.0px; padding: 1px; text-align: center;">
    <b>מטלה 3</b>
</div>


<div dir="rtl">
    
### מטרת המטלה הינה לתרגל את החומר שנלמד בשיעור ובתרגול השלישי

   - **שימו לב שהמחברת נבדקת אוטומטית - לכן עיבדו בדיוק לפי ההוראות**
   - **אין למחוק תאים**
   - **אין להוסיף תאים**
   - **אין להגיש אף מחברת אחרת מכל סוג שהוא - יש להגיש את מחברת זו בלבד** 
       - **תא הגשת תשובתכם לשאלה נמצא תחת השאלה בתא עם הערה # YOUR CODE HERE**
       - **מחקו את הקוד המצורף raise NotImplementedError(מחק.י אותי) והכניסו את תשובתכם**
#### בהצלחה !
</div>

<div style="color:#03112A; border-radius:8px; background-color:#d0e8ff; font-size:150%; letter-spacing:1.0px; padding: 1px; text-align: center;">
    <b>שאלה 1</b>
</div>

<div dir="rtl">

קראו את הקובץ הבא :
https://raw.githubusercontent.com/nlihin/data-analytics/main/datasets/aircraft%20wildlife%20strikes%202018-2020.csv

**לא להגיש את המטלה בקובץ שקראתם מהמחשב !**

עבור הקובץ שקראתם, קחו את שדה התאריך, וצרו ממנו עמודה חדשה בשם <code> weekday</code>
    <br>
    בעמודה זו צריך להיות את השם המלא של היום בשבוע בו התרחשה התאונה. לדוגמה:
    <br>
    <code> Monday, Tuesday,.., Sunday </code>

 לנוחיותכם : בתא הפתרון שלכם מצורף קריאת בצורה **הנכונה**
</div>

<!-- grade_id: cell-e91f30c44bb94ab1, points: 0, grade: false, solution: true -->
```python
import pandas as pd

df_path = "https://raw.githubusercontent.com/nlihin/data-analytics/main/datasets/aircraft%20wildlife%20strikes%202018-2020.csv"
df = pd.read_csv(df_path)

# YOUR CODE HERE
raise NotImplementedError(מחק.י אותי)
```

<!-- grade_id: cell-60e16e2163282990, points: 20, grade: true, solution: false -->
```python
# לא לכתוב או למחוק תא זה
```

<div style="color:#03112A; border-radius:8px; background-color:#d0e8ff; font-size:150%; letter-spacing:1.0px; padding: 1px; text-align: center;">
    <b>שאלה 2</b>
</div>

<div dir="rtl" style="text-align: right;">
באיזה יום בשבוע היו הכי הרבה דיווחים?
שימרו את תשובתכם בתוך משתנה בשם<br>
<code>top_day</code><br>
במשתנה ששמרתם יוצג היום בשבוע בלבד

<!-- grade_id: cell-ceb86cf2a2a16239, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError(מחק.י אותי)
```

<!-- grade_id: cell-40f75be14d0c85d0, points: 20, grade: true, solution: false -->
```python
# לא לכתוב או למחוק תא זה

```

<div style="color:#03112A; border-radius:8px; background-color:#d0e8ff; font-size:150%; letter-spacing:1.0px; padding: 1px; text-align: center;">
    <b>שאלה 3</b>
</div>

<div dir="rtl" style="text-align: right;">
מיצאו את חברת התעופה שאצלה התרחשו הכי הרבה תאונות עם ציפורים גדולות ביום שישי 
    <br>
    תוכלו למצוא את גודל הציפור בעמודה
    <code>SIZE </code>
    <br>
    ציפורים גדולות הן מסוג
    <code>Large </code>
    <br>
    כמה תאונות עם ציפורים גדולות התרחשו עבור חברה זו ביום זה?
    שימרו את תשובתכם במשתנה בשם
    <code>operator_friday_large_size_accidents</code>
    <br>
    שימו לב שחברת התעופה צריכה להיות בעלת שם כלשהו, לא Unknown
 </div>

<!-- grade_id: cell-27c67d14db38661a, points: 0, grade: false, solution: true -->
```python
# Solution cell

#operator_friday_large_size_accidents = 

# YOUR CODE HERE
raise NotImplementedError(מחק.י אותי)

```

<!-- grade_id: cell-49ac6ea7c7605257, points: 20, grade: true, solution: false -->
```python
# לא לכתוב או למחוק תא זה
```

<div style="color:#03112A; border-radius:8px; background-color:#d0e8ff; font-size:150%; letter-spacing:1.0px; padding: 1px; text-align: center;">
    <b>שאלה 4</b>
</div>

<div dir="rtl" style="text-align: right;">
    
אתם עובדים בטרמינל כמנתחי נתונים והתבקשתם למצוא **באיזה יום** הייתה טיסה שמקורה `SOURCE` הוא `Air Transport` והמפעיל `OPERATOR` שלו הוא לא ידוע `UNKNOWN`
שמרו את המשתנה בשם `Unknown_incident_operator`

**יש לשמור את ערך הפלט בלבד**

</div>

<!-- grade_id: cell-7c101949e141e226, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError(מחק.י אותי)
```

<!-- grade_id: cell-55e3b21b7c446251, points: 20, grade: true, solution: false -->
```python
# לא לכתוב או למחוק תא זה
# בדיקה אם המשתנה Unknown_incident_operator קיים
assert 'Unknown_incident_operator' in globals(), "The variable 'Unknown_incident_operator' should be defined."

# בדיקה אם המשתנה Unknown_incident_operator הוא מסוג str
assert isinstance(Unknown_incident_operator, str), "The variable 'Unknown_incident_operator' should be a string representing the day of the week."

# בדיקה אם הערך של Unknown_incident_operator נכון
expected_value = df[
    (df['SOURCE'].str.contains('Air Transport')) & 
    (df['OPERATOR'].str.contains('UNKNOWN'))
]['weekday'].values[0]
assert Unknown_incident_operator == expected_value, (
    f"The value of 'Unknown_incident_operator' is incorrect. Expected '{expected_value}', but got '{Unknown_incident_operator}'."
)

```

<div style="color:#03112A; border-radius:8px; background-color:#d0e8ff; font-size:150%; letter-spacing:1.0px; padding: 1px; text-align: center;">
    <b>שאלה 5</b>
</div>

<div dir="rtl" style="text-align: right;">

מהו מספר השורה של הטיסה שמצאתם בשאלה 4 ?

עשו שימוש ב- `iloc` להציג בשורה זו את הערך בעמודה `RUNWAY`

שמרו את המשתנה בשם `Runway_code`

**יש לשמור את ערך הפלט בלבד**

</div>

<!-- grade_id: cell-8795e8a3aa342b04, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError(מחק.י אותי)
```

<!-- grade_id: cell-29381b599ccbcf13, points: 20, grade: true, solution: false -->
```python
# לא לכתוב או למחוק תא זה

```

<div dir="rtl">

**אסור להוסיף תאים נוספים מתחת לתא זה!!**

**בהצלחה !**
</div>